# Tutorial 7: The Language of the System

__The goal of this assignment is to answer 10 questions related to [a talk on programming from Rich Hickey](https://youtu.be/ROor6_NGIWU).__

I think this talk is important to develop your critical regard on the way to architecture large processing systems.

I also hope this talk will open your mind on new programming constructs to support these architectures.

**For each question, your function should return the number associated to the answer.**

There is only one valid answer per question.

**Example: What is the sound of a dog ?**
1. meow
2. **woof**
3. hoot
4. honk

```python
def Q0():
    return 2
```

__Grade scale__: 10 points
- __correct answer__: 1 points
- __incorrect answer__: 0 points

```python
from IPython.display import HTML

HTML('<iframe width="980" height="600" src="https://www.youtube.com/embed/ROor6_NGIWU" frameborder="0"></iframe>')
```

Output:
```
<IPython.core.display.HTML object>
```

# Questions

__1. Why does a programmer write a program ?__
1. To affect the world outside of the program
2. To create a problem for someone else
3. To generate heat from a computer
4. To provide the answer to life

<!-- grade_id: cell-af420023b7676ec1, points: 0, grade: false, solution: true -->
```python
def Q1():
    ### BEGIN SOLUTION
    return 1
    ### END SOLUTION

Q1()
```

Output:
```
1
```

<!-- grade_id: cell-38b42e490724ca6d, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q1(), int)
### BEGIN HIDDEN TESTS
assert Q1() == 1
### END HIDDEN TESTS
```

__2. What is a system ?__
1. Language + Runtime
2. An ensemble of programs
3. A program that stores data
4. A program that manages resources

<!-- grade_id: cell-b969bb555ff6e41f, points: 0, grade: false, solution: true -->
```python
def Q2():
    ### BEGIN SOLUTION
    return 2
    ### END SOLUTION

Q3()
```

Output:
```
3
```

<!-- grade_id: cell-12c7bfa02b131299, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q2(), int)
### BEGIN HIDDEN TESTS
assert Q2() == 2
### END HIDDEN TESTS
```

__3. Which item is NOT an exchange format ?__
1. JSON
2. Jupyter Notebook
3. Protocol Buffer
4. CORBA

<!-- grade_id: cell-4a5f68774ae2d9f4, points: 0, grade: false, solution: true -->
```python
def Q3():
    ### BEGIN SOLUTION
    return 4
    ### END SOLUTION

Q3()
```

Output:
```
3
```

<!-- grade_id: cell-d267f272cd2eca06, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q3(), int)
### BEGIN HIDDEN TESTS
assert Q3() == 4
### END HIDDEN TESTS
```

__4. What is false about system names ?__
1. they are global
2. they are reversed for functions/verbs
3. they have a fully qualified namespace
4. we fight to reserve them (e.g. on github)

<!-- grade_id: cell-5151e01f60804944, points: 0, grade: false, solution: true -->
```python
def Q4():
    ### BEGIN SOLUTION
    return 2
    ### END SOLUTION

Q4()
```

Output:
```
4
```

<!-- grade_id: cell-481746201e851626, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q4(), int)
### BEGIN HIDDEN TESTS
assert Q4() == 2
### END HIDDEN TESTS
```

__5. Which property of a service is dangerous in the large ?__
1. correctness
2. statefulness
3. accomplish work
4. distributed processing

<!-- grade_id: cell-0d6b6981c043eae0, points: 0, grade: false, solution: true -->
```python
def Q5():
    ### BEGIN SOLUTION
    return 2
    ### END SOLUTION

Q5()
```

Output:
```
4
```

<!-- grade_id: cell-d7fc18b7ec8a407d, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q5(), int)
### BEGIN HIDDEN TESTS
assert Q5() == 2
### END HIDDEN TESTS
```

__6. What method should be avoided on values in the large ?__
1. transform
2. record
3. route
4. update

<!-- grade_id: cell-ed5f0b03a1675a58, points: 0, grade: false, solution: true -->
```python
def Q6():
    ### BEGIN SOLUTION
    return 4
    ### END SOLUTION

Q6()
```

Output:
```
3
```

<!-- grade_id: cell-93ef4ae7112f9795, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q6(), int)
### BEGIN HIDDEN TESTS
assert Q6() == 4
### END HIDDEN TESTS
```

__7. Which data structure should be used to move values between programs ?__
1. sets
2. lists
3. stacks
4. queues

<!-- grade_id: cell-6a1daef61811752d, points: 0, grade: false, solution: true -->
```python
def Q7():
    ### BEGIN SOLUTION
    return 4
    ### END SOLUTION

Q7()
```

Output:
```
1
```

<!-- grade_id: cell-fe9ea22acb67684f, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q7(), int)
### BEGIN HIDDEN TESTS
assert Q7() == 4
### END HIDDEN TESTS
```

__8. Which concept is NOT important in the Epochal Time Model ?__
1. state
2. process
3. references
4. databases

<!-- grade_id: cell-803731b2fe93baa1, points: 0, grade: false, solution: true -->
```python
def Q8():
    ### BEGIN SOLUTION
    return 4
    ### END SOLUTION

Q8()
```

Output:
```
1
```

<!-- grade_id: cell-3916000bef44802c, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q8(), int)
### BEGIN HIDDEN TESTS
assert Q8() == 4
### END HIDDEN TESTS
```

__9. What is true about system failures ?__
1. we can correlate failures across services
2. we can guarantee that another service is online
3. we can guarantee that something bad will happen
4. we can wait for a system to recover from a failure

<!-- grade_id: cell-c51ba59356d3f7f0, points: 0, grade: false, solution: true -->
```python
def Q9():
    ### BEGIN SOLUTION
    return 3
    ### END SOLUTION

Q9()
```

Output:
```
1
```

<!-- grade_id: cell-7a48f41b09b209a3, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q9(), int)
### BEGIN HIDDEN TESTS
assert Q9() == 3
### END HIDDEN TESTS
```

__10. What should NOT be part of a language for the system ?__
1. a fix communication interface
2. applications as services
3. protocols and formats
4. message passing

<!-- grade_id: cell-c5c2bdae3fd3601f, points: 0, grade: false, solution: true -->
```python
def Q10():
    ### BEGIN SOLUTION
    return 1
    ### END SOLUTION

Q10()
```

Output:
```
3
```

<!-- grade_id: cell-3a108cbb112130ff, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q10(), int)
### BEGIN HIDDEN TESTS
assert Q10() == 1
### END HIDDEN TESTS
```